### COdio Activity 2.7: Approximating Correlation from scatterplots

**Expected Time: 30 Minutes**

**Total Points: 25**

This assignment aims to reinforce your understanding of correlation.  A scatterplot of two random variables is shown in an image, and you are asked to guess the correlation between these variables.  Your guesses should be with $0.3$ of the actual correlation value.

## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)
- [Problem 5](#Problem-5)


```python
import scipy.stats as stats
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
import numpy as np
```

[Back to top](#Index:) 

### Problem 1

**5 Points**

Given the scatterplot below, guess the correlation between the features `x1` and `y`. Assign your answer as a float to `ans1` below.  An answer within .3 of the true correlation will receive credit.

![](images/guess_1.png)

<!-- grade_id: cell-d8bbfa13f2b7ae0a, points: 0, grade: false, solution: true -->
```python
### GRADED

ans1 = None

### BEGIN SOLUTION
ans1 = 0.56763988
### END SOLUTION

# Answer check
print(ans1)
print(type(ans1))
```

Output:
```
0.56763988
<class 'float'>
```

<!-- grade_id: cell-2173092ece0a4682, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans1_ = 0.56763988
#
#
#
assert abs(ans1_ - ans1) < 0.3
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2

**5 Points**


Given the scatterplot below, guess the correlation between the features `x2` and `y`. Assign your answer as a float to `ans2` below.  An answer within .3 of the true correlation will receive credit.

![](images/guess_2.png)

<!-- grade_id: cell-5d7037f91eeb36be, points: 0, grade: false, solution: true -->
```python
### GRADED

ans2 = None

### BEGIN SOLUTION
ans2 = 0.11112693
### END SOLUTION

# Answer check
print(ans2)
print(type(ans2))
```

Output:
```
0.11112693
<class 'float'>
```

<!-- grade_id: cell-b1b7fa4feb4f789f, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans2_ = 0.11112693
#
#
#
assert abs(ans2_ - ans2) < 0.3
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 3

**5 Points**

Given the scatterplot below, guess the correlation between the features `x3` and `y`. Assign your answer as a float to `ans3` below.  An answer within .3 of the true correlation will receive credit.

![](images/guess_3.png)

<!-- grade_id: cell-54c7bc870388a32e, points: 0, grade: false, solution: true -->
```python
### GRADED

ans3 = None

### BEGIN SOLUTION
ans3 = 0.79386772
### END SOLUTION

# Answer check
print(ans3)
print(type(ans3))
```

Output:
```
0.79386772
<class 'float'>
```

<!-- grade_id: cell-ca9169fc61fd4941, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans3_ = 0.79386772
#
#
#
assert abs(ans3_ - ans3) < 0.3
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4

**5 Points**


Given the scatterplot below, guess the correlation between the features `x4` and `y`. Assign your answer as a float to `ans4` below.  An answer within .3 of the true correlation will receive credit.

![](images/guess_4.png)

<!-- grade_id: cell-d337632c05c8c07e, points: 0, grade: false, solution: true -->
```python
### GRADED

ans4 = None

### BEGIN SOLUTION
ans4 = -0.3644958
### END SOLUTION

# Answer check
print(ans4)
print(type(ans4))
```

Output:
```
-0.3644958
<class 'float'>
```

<!-- grade_id: cell-094f8ab7c512e017, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans4_ = -0.3644958
#
#
#
assert abs(ans4_ - ans4) < 0.3
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 5

**5 Points**

Given the scatterplot below, guess the correlation between the features `x5` and `y`. Assign your answer as a float to `ans5` below.  An answer within .3 of the true correlation will receive credit.

![](images/guess_5.png)

<!-- grade_id: cell-f21cf8efefa20152, points: 0, grade: false, solution: true -->
```python
### GRADED

ans5 = None

### BEGIN SOLUTION
ans5 = -0.09696047
### END SOLUTION

# Answer check
print(ans5)
print(type(ans5))
```

Output:
```
-0.09696047
<class 'float'>
```

<!-- grade_id: cell-19b36c42fe499d50, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans5_ = -0.09696047
#
#
#
assert abs(ans5_ - ans5) < 0.3
### END HIDDEN TESTS
```