### Codio Activity 2.1: Continuous Uniform Distributions

**Expected Time**: 45 Minutes

**Total Points**: 10

This activity focuses on creating and visualizing continuous probability distributions using `scipy.stats`. You will get practice creating uniform distribution objects, examining their summary statistics (mean, variance, standard deviation), and visualizing the distribution using `matplotlib`. 

#### Index:

- [Problem 1](#Problem-1:-Creating-a-Distribution-Object)
- [Problem 2](#Problem-2:-Descriptive-Statistics-of-Distribution)
- [Problem 3](#Problem-3:-Plotting-the-Distribution)
- [Problem 4](#Problem-4:-Evaluating-Probabilities)
- [Problem 5](#Problem-5:-Cumulative-Probabilities)

```python
from scipy.stats import uniform
from scipy.stats import norm
import matplotlib.pyplot as plt
import numpy as np
```

[Back to top](#Index:) 

### Problem 1: Creating a Distribution Object

**2 Points**

To begin, use the above imported `uniform` function to create a uniform distribution object with `loc = 10` and `scale = 3`.  Assign your solution as a distribution object to `dist1` below.

<!-- grade_id: cell-49b8f024d62ed17a, points: 0, grade: false, solution: true -->
```python
### GRADED

dist1 = ''

### BEGIN SOLUTION
dist1 = uniform(loc = 10, scale = 3)
### END SOLUTION

# Answer check
print(type(dist1))
```

Output:
```
<class 'scipy.stats._distn_infrastructure.rv_frozen'>
```

<!-- grade_id: cell-504b7bdf18ac95c5, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
dist1_ = uniform(loc = 10, scale = 3)
#
#
#
assert type(dist1_) == type(dist1)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2: Descriptive Statistics of Distribution

**2 Points**

Using your distribution object from problem 1, determine the mean, variance, and standard deviation of the distribution using its own built in methods.  Assign these as floats to the variables, `dist1_mean`, `dist1_var`, and `dist1_std` below.

<!-- grade_id: cell-f9b4f824fd625222, points: 0, grade: false, solution: true -->
```python
### GRADED

dist1_mean = ''
dist1_var = ''
dist1_std = ''

### BEGIN SOLUTION
dist1 = uniform(loc = 10, scale = 3)
dist1_mean = dist1.mean()
dist1_var = dist1.var()
dist1_std = dist1.std()
### END SOLUTION

# Answer check
print([dist1_mean, dist1_var, dist1_std])
```

Output:
```
[11.5, 0.75, 0.8660254037844386]
```

<!-- grade_id: cell-b4ced79006d31e07, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
dist1_ = uniform(loc = 10, scale = 3)
dist1_mean_ = dist1_.mean()
dist1_var_ = dist1_.var()
dist1_std_ = dist1_.std()
#
#
#
assert dist1_mean == dist1_mean_, 'Means are not equal'
assert dist1_var == dist1_var_, 'Variance is not equal.'
assert dist1_std == dist1_std_, 'Standard Deviation is not equal'
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 3: Plotting the Distribution                                                              

**2 Points**

Note that the distribution is bound by `loc` and `loc + scale` on the left and right sides respectively.  To build a plot, we will give one unit extra on each so we see the entire distribution.  To do so, define `x` using `np.linspace` and create an array of 100 evenly spaced values starting at 9 and ending at 14.  Then, use these `x` values with `dist1`'s `.pdf` method to produce a plot of the distribution using `plt.plot`.  You should have a plot similar to that shown below.



![](images/dist1.png)









<!-- grade_id: cell-55be2c1f62142ff0, points: 0, grade: false, solution: true -->
```python
### GRADED

x = ''

### BEGIN SOLUTION
x = np.linspace(9, 14, 100)
### END SOLUTION

# Answer check
print(min(x), max(x))
```

Output:
```
9.0 14.0
```

<!-- grade_id: cell-f20857cdf84a3f91, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
x_ = np.linspace(9, 14, 100)
#
#
#
np.testing.assert_array_equal(x, x_, err_msg='Make sure 100 points from 9 to 14 with linspace')
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4: Evaluating Probabilities 

**2 Points**

Finally, you want to use the `.pdf` method to evaluate the probability of a specific outcome.  Determine the following probabilities using `dist1` and assign to the respective variables as floats below.

- $P(9)$
- $P(11)$
- $P(11~ \text{or} ~12)$

<!-- grade_id: cell-1f288c9a16d0b2a5, points: 0, grade: false, solution: true -->
```python
### GRADED

p_9 = ''
p_11 = ''
p_11_or_12 = ''

### BEGIN SOLUTION
p_9 = dist1.pdf(9)
p_11 = dist1.pdf(11)
p_11_or_12 = dist1.pdf(11) + dist1.pdf(12)
### END SOLUTION

# Answer check
print(p_9, p_11, p_11_or_12)
```

Output:
```
0.0 0.3333333333333333 0.6666666666666666
```

<!-- grade_id: cell-d964dbefc2c842b6, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
dist1_ = uniform(loc = 10, scale = 3)
p_9_ = dist1_.pdf(9)
p_11_ = dist1_.pdf(11)
p_11_or_12_ = dist1_.pdf(11) + dist1_.pdf(12)
#
#
#
assert p_9 == p_9_, 'Check P(9)'
assert p_11_ == p_11, 'Check P(11)'
assert p_11_or_12 == p_11_or_12_, 'Check P(11 or 12)'
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 5: Cumulative Probabilities

**2 Points**

The `.cdf` method can be used to determine cumulative probabilities.  Use the `.cdf` method of `dist1` to determine the following probabilities and assign them as a float to the relevant variable below. 

- $P(x < 12)$
- $P(11 < x < 13)$
- $P(x > 12)$

<!-- grade_id: cell-6a64fed590ac20b9, points: 0, grade: false, solution: true -->
```python
### GRADED

p_less_than_12 = ''
p_between_11_and_13 = ''
p_greater_than_12 = ''

### BEGIN SOLUTION
p_less_than_12 = dist1.cdf(12)
p_between_11_and_13 = dist1.cdf(13) - dist1.cdf(11)
p_greater_than_12 = 1 - dist1.cdf(12)
### END SOLUTION

# Answer check
print(p_less_than_12, p_between_11_and_13, p_greater_than_12)
```

Output:
```
0.6666666666666666 0.6666666666666667 0.33333333333333337
```

<!-- grade_id: cell-e85736fbeaca8790, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
p_less_than_12_ = dist1_.cdf(12)
p_between_11_and_13_ = dist1_.cdf(13) - dist1_.cdf(11)
p_greater_than_12_ = 1 - dist1_.cdf(12)
#
#
#
assert p_less_than_12 == p_less_than_12_, 'Check P(x < 12)'
assert p_between_11_and_13 == p_between_11_and_13_, 'Check P(11 < x < 13)'
assert p_greater_than_12 == p_greater_than_12_, 'Check P(x > 12)'
### END HIDDEN TESTS
```