## Codio Activity 2.6: Computing the Sample Covariance Matrix

**Expected Time: 45 Minutes**

**Total Points: 20**

This activity focuses on using covariance and correlation to understand the strength of relationships between two quantitative variables.  Also, you will use a heatmap to visualize these relationships.  Finally, you identify situations where correlation is not an appropriate measure based on the assumption of a linear variable.

## Index:

- [Problem 1](#Problem-1:-Calculating-Covariance)
- [Problem 2](#Problem-2:-Calculating-Correlation)
- [Problem 3](#Problem-3:-Heatmap-of-Covariance)
- [Problem 4](#Problem-4:-Heatmap-of-Correlation)
- [Problem 5](#Problem-5:-Strongest-Positive-Correlation)
- [Problem 6](#Problem-6:-Interpreting-negative-correlation)


### Covariance and Correlation

In probability theory and statistics, the mathematical concepts of covariance and correlation are very similar. Both describe the degree to which two random variables or sets of random variables tend to deviate from their expected values in similar ways.

If X and Y are two random variables, with means (expected values)  $\mu_X$ and  $\mu_Y$ and standard deviations  $σ_X$ and  $σ_Y$, respectively, then their covariance and correlation are as follows:

covariance $$\text{cov}_{XY}=\sigma_{XY}=E[(X-\mu_{X})\,(Y-\mu_{Y})]$$

correlation $$\text{corr}_{XY} = \rho_{XY} = E[(X - \mu_X)(Y - \mu_Y)]/(\sigma_X \sigma_Y)$$

so that

$$\rho_{XY}=\sigma_{XY}/(\sigma_{X} \sigma_{Y})$$


where E is the expected value operator. Notably, correlation is dimensionless while covariance is in units obtained by multiplying the units of the two variables. --[source](https://en.wikipedia.org/wiki/Covariance_and_correlation)

#### Dataset: Penguins

Our sample dataset comes from the `seaborn` visualization library and contains measurements of penguins of different species located on different islands.  Below, we load the data and display the first few rows and datatype information.

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import seaborn as sns
```

```python
penguin_df = pd.read_csv('data/penguins.csv')
```

```python
penguin_df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 333 entries, 0 to 332
Data columns (total 7 columns):
species              333 non-null object
island               333 non-null object
bill_length_mm       333 non-null float64
bill_depth_mm        333 non-null float64
flipper_length_mm    333 non-null float64
body_mass_g          333 non-null float64
sex                  333 non-null object
dtypes: float64(4), object(3)
memory usage: 18.3+ KB
```

```python
penguin_df.head()
```

Output:
```
species     island  bill_length_mm  bill_depth_mm  flipper_length_mm  \
0  Adelie  Torgersen            39.1           18.7              181.0   
1  Adelie  Torgersen            39.5           17.4              186.0   
2  Adelie  Torgersen            40.3           18.0              195.0   
3  Adelie  Torgersen            36.7           19.3              193.0   
4  Adelie  Torgersen            39.3           20.6              190.0   

   body_mass_g     sex  
0       3750.0    Male  
1       3800.0  Female  
2       3250.0  Female  
3       3450.0  Female  
4       3650.0    Male
```

[Back to top](#Index:) 

### Problem 1: Calculating Covariance

**4 Points**

Pandas DataFrame's have a built in `.cov()` method that will return the covariation between each of the numeric features.  Use the `.cov()` method below to create the covariance matrix.  Assign your solution to the variable `penguins_cov` below.

<!-- grade_id: cell-cb793d1142c2d0b8, points: 0, grade: false, solution: true -->
```python
### GRADED

penguins_cov = ''

### BEGIN SOLUTION
penguins_cov = penguin_df.cov()
### END SOLUTION
```

<!-- grade_id: cell-52236f3f94987fe5, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
penguin_df_ = pd.read_csv('data/penguins.csv')
penguins_cov_ = penguin_df_.cov()
pd.testing.assert_frame_equal(penguins_cov, penguins_cov_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2: Calculating Correlation

**4 Points**

Similarly, pandas DataFrames have a built in `.corr()` method that computes the correlation matrix of the numeric columns.  Below, use the `.corr()` method on the `penguins_df` and assign your results to the variable `penguin_corr` below.

<!-- grade_id: cell-bdd62ffdd36dbed4, points: 0, grade: false, solution: true -->
```python
### GRADED

penguins_corr = ''

### BEGIN SOLUTION
penguins_corr = penguin_df.corr()
### END SOLUTION
```

<!-- grade_id: cell-9e1459665935d821, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
penguin_df_ = pd.read_csv('data/penguins.csv')
penguins_corr_ = penguin_df_.corr()
pd.testing.assert_frame_equal(penguins_corr, penguins_corr_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 3: Heatmap of Covariance

**0 Points**

The covariance matrix can be visualized to color according to the strength of associations.  While `matplotlib` has a `.pcolor()` function that can do this, the `seaborn` heatmap is better suited for this visualization.  Use the `sns.heatmap()` function to visualize the covariance matrix.  Set the `annot` keyword argument to `True` so that the covariance values are displayed in the heatmap.

<!-- grade_id: cell-3a797a4238076ff2, points: 0, grade: false, solution: true -->
```python
### GRADED

### BEGIN SOLUTION
pass
### END SOLUTION
```

<!-- grade_id: cell-c070e03d523c773a, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert 1 == 1
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4: Heatmap of Correlation

**0 Points**

Again using `sns.heatmap()`, visualize the correlation matrix for the penguins data.  Set the `annot` argument to `True` again.

<!-- grade_id: cell-382d431a32a827d7, points: 0, grade: false, solution: true -->
```python
### GRADED

### BEGIN SOLUTION
pass
### END SOLUTION
```

<!-- grade_id: cell-fca610fa8f05bb27, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert 1 == 1, 'This is not graded.'
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 5: Strongest Positive Correlation

**8 Points**

What two features have the strongest positive correlation?  Enter your answer as a list of strings using the exact column names to `highest_pos_corr` below.

<!-- grade_id: cell-c1c0856c81e9ba49, points: 0, grade: false, solution: true -->
```python
### GRADED

highest_pos_corr = ''

### BEGIN SOLUTION
highest_pos_corr = ['body_mass_g', 'flipper_length_mm']
### END SOLUTION
```

<!-- grade_id: cell-147fd8a9ba9c4d9b, points: 8, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
highest_pos_corr_ = ['body_mass_g', 'flipper_length_mm']
assert set(highest_pos_corr_) == set(highest_pos_corr)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 6: Interpreting negative correlation

**4 Points**

The correlation between `body_mass_g` and `bill_depth_mm` reads -0.47.  Complete the following sentence that interprets this correlation.

```
As body mass increases bill depth ___________
```

Assign your solution as a string -- either 'increases', 'decreases', or 'neither' to `ans_6` below.

<!-- grade_id: cell-b68d1118f1213e2a, points: 0, grade: false, solution: true -->
```python
### GRADED

ans_6 = ''

### BEGIN SOLUTION
ans_6 = 'decreases'
### END SOLUTION
```

<!-- grade_id: cell-64112c10a27b811b, points: 4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
ans_6_ = 'decreases'
assert ans_6 == ans_6_
### END HIDDEN TESTS
```